[1604159804] SERVICE ALERT: localhost;HTTP;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 80: Connection refused
[1604159864] SERVICE ALERT: localhost;HTTP;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 80: Connection refused
[1604159924] SERVICE ALERT: localhost;HTTP;CRITICAL;SOFT;3;connect to address 127.0.0.1 and port 80: Connection refused
[1604159984] SERVICE ALERT: localhost;HTTP;CRITICAL;HARD;4;connect to address 127.0.0.1 and port 80: Connection refused
[1604178466] HOST ALERT: backup;DOWN;SOFT;1;(Return code of 13 for host 'backup' was out of bounds)
[1604178495] SERVICE ALERT: backup;CPU;CRITICAL;HARD;1;(Return code of 13 for service 'CPU' on host 'backup' was out of bounds)
[1604178502] SERVICE ALERT: elk;CPU;CRITICAL;SOFT;1;(Return code of 13 for service 'CPU' on host 'elk' was out of bounds)
[1604178502] HOST ALERT: elk;DOWN;SOFT;1;(Return code of 13 for host 'elk' was out of bounds)
[1604178509] SERVICE ALERT: web;CPU;CRITICAL;SOFT;1;(Return code of 13 for service 'CPU' on host 'web' was out of bounds)
[1604178509] HOST ALERT: web;DOWN;SOFT;1;(Return code of 13 for host 'web' was out of bounds)
[1604178524] SERVICE ALERT: backup;Current Load;CRITICAL;HARD;1;(Return code of 13 for service 'Current Load' on host 'backup' was out of bounds)
[1604178526] HOST ALERT: backup;DOWN;SOFT;2;(Return code of 13 for host 'backup' was out of bounds)
[1604178531] SERVICE ALERT: elk;Current Load;CRITICAL;HARD;1;(Return code of 13 for service 'Current Load' on host 'elk' was out of bounds)
[1604178539] SERVICE ALERT: web;Current Load;CRITICAL;HARD;1;(Return code of 13 for service 'Current Load' on host 'web' was out of bounds)
[1604178553] SERVICE ALERT: backup;Current Users;CRITICAL;HARD;1;(Return code of 13 for service 'Current Users' on host 'backup' was out of bounds)
[1604178561] SERVICE ALERT: elk;Current Users;CRITICAL;HARD;1;(Return code of 13 for service 'Current Users' on host 'elk' was out of bounds)
[1604178562] SERVICE ALERT: elk;CPU;CRITICAL;HARD;4;(Return code of 13 for service 'CPU' on host 'elk' was out of bounds)
[1604178562] HOST ALERT: elk;DOWN;SOFT;2;(Return code of 13 for host 'elk' was out of bounds)
[1604178568] SERVICE ALERT: web;Current Users;CRITICAL;HARD;1;(Return code of 13 for service 'Current Users' on host 'web' was out of bounds)
[1604178569] SERVICE ALERT: web;CPU;CRITICAL;HARD;4;(Return code of 13 for service 'CPU' on host 'web' was out of bounds)
[1604178569] HOST ALERT: web;DOWN;SOFT;2;(Return code of 13 for host 'web' was out of bounds)
[1604178583] SERVICE ALERT: backup;HTTP;CRITICAL;HARD;1;(Return code of 13 for service 'HTTP' on host 'backup' was out of bounds)
[1604178586] HOST ALERT: backup;DOWN;SOFT;3;(Return code of 13 for host 'backup' was out of bounds)
[1604178590] SERVICE ALERT: elk;HTTP;CRITICAL;HARD;1;(Return code of 13 for service 'HTTP' on host 'elk' was out of bounds)
[1604178597] SERVICE ALERT: web;HTTP;CRITICAL;HARD;1;(Return code of 13 for service 'HTTP' on host 'web' was out of bounds)
[1604178612] SERVICE ALERT: backup;PING;CRITICAL;HARD;1;(Return code of 13 for service 'PING' on host 'backup' was out of bounds)
[1604178619] SERVICE ALERT: elk;PING;CRITICAL;HARD;1;(Return code of 13 for service 'PING' on host 'elk' was out of bounds)
[1604178622] HOST ALERT: elk;DOWN;SOFT;3;(Return code of 13 for host 'elk' was out of bounds)
[1604178626] SERVICE ALERT: web;PING;CRITICAL;HARD;1;(Return code of 13 for service 'PING' on host 'web' was out of bounds)
[1604178629] HOST ALERT: web;DOWN;SOFT;3;(Return code of 13 for host 'web' was out of bounds)
[1604178641] SERVICE ALERT: backup;RAM;CRITICAL;HARD;1;(Return code of 13 for service 'RAM' on host 'backup' was out of bounds)
[1604178646] HOST ALERT: backup;DOWN;SOFT;4;(Return code of 13 for host 'backup' was out of bounds)
[1604178648] SERVICE ALERT: elk;RAM;CRITICAL;HARD;1;(Return code of 13 for service 'RAM' on host 'elk' was out of bounds)
[1604178656] SERVICE ALERT: web;RAM;CRITICAL;HARD;1;(Return code of 13 for service 'RAM' on host 'web' was out of bounds)
[1604178670] SERVICE ALERT: backup;Root Partition;CRITICAL;HARD;1;(Return code of 13 for service 'Root Partition' on host 'backup' was out of bounds)
[1604178678] SERVICE ALERT: elk;Root Partition;CRITICAL;HARD;1;(Return code of 13 for service 'Root Partition' on host 'elk' was out of bounds)
[1604178682] HOST ALERT: elk;DOWN;SOFT;4;(Return code of 13 for host 'elk' was out of bounds)
[1604178685] SERVICE ALERT: web;Root Partition;CRITICAL;HARD;1;(Return code of 13 for service 'Root Partition' on host 'web' was out of bounds)
[1604178689] HOST ALERT: web;DOWN;SOFT;4;(Return code of 13 for host 'web' was out of bounds)
[1604178700] SERVICE ALERT: backup;SSH;CRITICAL;HARD;1;(Return code of 13 for service 'SSH' on host 'backup' was out of bounds)
[1604178706] HOST ALERT: backup;DOWN;SOFT;5;(Return code of 13 for host 'backup' was out of bounds)
[1604178707] SERVICE ALERT: elk;SSH;CRITICAL;HARD;1;(Return code of 13 for service 'SSH' on host 'elk' was out of bounds)
[1604178714] SERVICE ALERT: web;SSH;CRITICAL;HARD;1;(Return code of 13 for service 'SSH' on host 'web' was out of bounds)
[1604178729] SERVICE ALERT: backup;Swap Usage;CRITICAL;HARD;1;(Return code of 13 for service 'Swap Usage' on host 'backup' was out of bounds)
[1604178736] SERVICE ALERT: elk;Swap Usage;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: Connection refused
[1604178746] HOST ALERT: elk;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.81 ms
[1604178748] HOST ALERT: web;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.12 ms
[1604178763] HOST ALERT: backup;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.70 ms
[1604178766] SERVICE ALERT: elk;Zombie Processes;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: Connection refused
[1604178766] SERVICE ALERT: elk;Total Processes;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: Connection refused
[1604178768] SERVICE ALERT: localhost;Current Load;OK;HARD;4;OK - load average: 0.01, 0.06, 0.24
[1604178772] HOST ALERT: localhost;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.11 ms
[1604178806] SERVICE ALERT: localhost;Current Users;OK;HARD;4;USERS OK - 1 users currently logged in
[1604178824] SERVICE ALERT: backup;Current Load;OK;HARD;4;OK - load average: 0.06, 0.04, 0.05
[1604178826] SERVICE ALERT: elk;Total Processes;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: Connection refused
[1604178826] SERVICE ALERT: elk;Zombie Processes;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: Connection refused
[1604178839] SERVICE ALERT: web;Current Load;OK;HARD;4;OK - load average: 0.00, 0.01, 0.05
[1604178853] SERVICE ALERT: backup;Current Users;OK;HARD;4;USERS OK - 0 users currently logged in
[1604178855] SERVICE ALERT: backup;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 1.27 ms
[1604178868] SERVICE ALERT: web;Current Users;OK;HARD;4;USERS OK - 0 users currently logged in
[1604178869] SERVICE ALERT: web;CPU;UNKNOWN;HARD;4;NRPE: Unable to read output
[1604178885] SERVICE ALERT: localhost;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.10 ms
[1604178886] SERVICE ALERT: elk;Total Processes;CRITICAL;SOFT;3;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: Connection refused
[1604178886] SERVICE ALERT: elk;Zombie Processes;CRITICAL;SOFT;3;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: Connection refused
[1604178903] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.050 second response time
[1604178918] SERVICE ALERT: localhost;Root Partition;OK;HARD;4;DISK OK - free space: / 36826 MB (89.95% inode=100%):
[1604178930] SERVICE ALERT: web;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 1.46 ms
[1604178941] SERVICE ALERT: backup;RAM;UNKNOWN;HARD;4;NRPE: Unable to read output
[1604178946] SERVICE ALERT: elk;Total Processes;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: Connection refused
[1604178946] SERVICE ALERT: elk;Zombie Processes;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: Connection refused
[1604178953] SERVICE ALERT: elk;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.94 ms
[1604178956] SERVICE ALERT: localhost;SSH;OK;HARD;4;SSH OK - OpenSSH_7.4 (protocol 2.0)
[1604178956] SERVICE ALERT: web;RAM;UNKNOWN;HARD;4;NRPE: Unable to read output
[1604178970] SERVICE ALERT: backup;Root Partition;OK;HARD;4;DISK OK - free space: / 37054 MiB (90.51% inode=100%):
[1604178985] SERVICE ALERT: web;Root Partition;OK;HARD;4;DISK OK - free space: / 37062 MiB (90.53% inode=100%):
[1604178993] SERVICE ALERT: localhost;Swap Usage;OK;HARD;4;SWAP OK - 99% free (2014 MB out of 2047 MB)
[1604179000] SERVICE ALERT: backup;SSH;OK;HARD;4;SSH OK - OpenSSH_7.4 (protocol 2.0)
[1604179007] SERVICE ALERT: elk;SSH;OK;HARD;4;SSH OK - OpenSSH_7.4 (protocol 2.0)
[1604179014] SERVICE ALERT: web;SSH;OK;HARD;4;SSH OK - OpenSSH_7.4 (protocol 2.0)
[1604179029] SERVICE ALERT: backup;Swap Usage;OK;HARD;4;SWAP OK - 99% free (2014 MB out of 2047 MB)
[1604179031] SERVICE ALERT: localhost;Total Processes;OK;HARD;4;PROCS OK: 39 processes with STATE = RSZDT
[1604181051] SERVICE ALERT: elk;RAM;UNKNOWN;HARD;4;NRPE: Unable to read output
[1604181052] SERVICE ALERT: elk;Zombie Processes;OK;HARD;4;PROCS OK: 0 processes with STATE = Z
[1604181052] SERVICE ALERT: elk;Total Processes;OK;HARD;4;PROCS OK: 107 processes
[1604181082] SERVICE ALERT: elk;Root Partition;OK;HARD;4;DISK OK - free space: / 34166 MiB (83.45% inode=99%):
[1604181141] SERVICE ALERT: elk;Swap Usage;OK;HARD;4;SWAP OK - 91% free (1843 MB out of 2047 MB)
[1604181265] SERVICE ALERT: elk;CPU;UNKNOWN;HARD;4;NRPE: Unable to read output
[1604181265] SERVICE ALERT: elk;Current Users;OK;HARD;4;USERS OK - 0 users currently logged in
[1604188709] SERVICE ALERT: elk;Current Load;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604188712] HOST ALERT: elk;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.1.242)
[1604188714] SERVICE ALERT: web;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188741] SERVICE ALERT: elk;Current Users;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604188747] SERVICE ALERT: web;Current Load;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188772] SERVICE ALERT: elk;Current Load;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604188772] SERVICE ALERT: elk;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.1.242)
[1604188772] HOST ALERT: elk;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.1.242)
[1604188774] SERVICE ALERT: web;CPU;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188779] SERVICE ALERT: web;Current Users;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused