[1604188806] SERVICE ALERT: elk;RAM;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604188807] SERVICE ALERT: web;Current Load;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188835] SERVICE ALERT: web;CPU;CRITICAL;SOFT;3;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188835] HOST ALERT: elk;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.1.242)
[1604188838] SERVICE ALERT: elk;Root Partition;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604188839] SERVICE ALERT: web;Current Users;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188844] SERVICE ALERT: web;RAM;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188868] SERVICE ALERT: web;Current Load;CRITICAL;SOFT;3;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188871] SERVICE ALERT: elk;SSH;CRITICAL;HARD;1;connect to address 192.168.1.242 and port 22: No route to host
[1604188876] SERVICE ALERT: web;Root Partition;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188895] SERVICE ALERT: web;CPU;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188898] HOST ALERT: elk;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.1.242)
[1604188899] SERVICE ALERT: web;Current Users;CRITICAL;SOFT;3;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188902] SERVICE ALERT: elk;Swap Usage;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604188904] SERVICE ALERT: web;RAM;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188928] SERVICE ALERT: web;Current Load;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188936] SERVICE ALERT: elk;Total Processes;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604188937] SERVICE ALERT: web;Root Partition;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188941] SERVICE ALERT: web;Swap Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188959] SERVICE ALERT: web;Current Users;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188962] HOST ALERT: elk;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.1.242)
[1604188964] SERVICE ALERT: web;RAM;CRITICAL;SOFT;3;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188967] SERVICE ALERT: web;Total Processes;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188967] SERVICE ALERT: web;Zombie Processes;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604188969] SERVICE ALERT: elk;CPU;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604188969] SERVICE ALERT: elk;Zombie Processes;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604188997] SERVICE ALERT: web;Root Partition;CRITICAL;SOFT;3;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604189001] SERVICE ALERT: web;Swap Usage;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604189024] SERVICE ALERT: web;RAM;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604189025] HOST ALERT: elk;DOWN;SOFT;6;CRITICAL - Host Unreachable (192.168.1.242)
[1604189027] SERVICE ALERT: web;Total Processes;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604189027] SERVICE ALERT: web;Zombie Processes;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604189057] SERVICE ALERT: web;Root Partition;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604189061] SERVICE ALERT: web;Swap Usage;CRITICAL;SOFT;3;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604189087] SERVICE ALERT: web;Total Processes;CRITICAL;SOFT;3;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604189087] SERVICE ALERT: web;Zombie Processes;CRITICAL;SOFT;3;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604189088] HOST ALERT: elk;DOWN;SOFT;7;CRITICAL - Host Unreachable (192.168.1.242)
[1604189122] SERVICE ALERT: web;Swap Usage;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604189147] SERVICE ALERT: web;Total Processes;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604189147] SERVICE ALERT: web;Zombie Processes;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.1.240 port 5666: Connection refused
[1604189152] HOST ALERT: elk;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.54 ms
[1604189169] SERVICE ALERT: elk;SSH;OK;SOFT;1;SSH OK - OpenSSH_7.4 (protocol 2.0)
[1604189376] SERVICE ALERT: elk;PING;OK;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.87 ms
[1604190081] SERVICE ALERT: elk;RAM;OK;HARD;4;RAM OK: 49% (116 of 235 MB) used
[1604190103] SERVICE ALERT: elk;Swap Usage;OK;HARD;4;SWAP OK - 95% free (1935 MB out of 2047 MB)
[1604190134] SERVICE ALERT: elk;Total Processes;OK;HARD;4;PROCS OK: 96 processes
[1604190167] SERVICE ALERT: elk;Zombie Processes;OK;HARD;4;PROCS OK: 0 processes with STATE = Z
[1604190270] SERVICE ALERT: elk;Current Users;OK;HARD;4;USERS OK - 1 users currently logged in
[1604190349] SERVICE ALERT: elk;Root Partition;OK;HARD;4;DISK OK - free space: / 34330 MiB (83.85% inode=99%):
[1604222493] SERVICE ALERT: backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 0.11, 0.34, 0.18
[1604222507] HOST ALERT: elk;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.1.242)
[1604222535] SERVICE ALERT: elk;Current Users;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604222554] SERVICE ALERT: backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.04, 0.28, 0.17
[1604222567] SERVICE ALERT: elk;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.1.242)
[1604222570] HOST ALERT: elk;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.1.242)
[1604222600] HOST ALERT: elk;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.53 ms
[1604222614] SERVICE ALERT: backup;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 0.11, 0.26, 0.17
[1604222674] SERVICE ALERT: backup;Current Load;WARNING;HARD;4;WARNING - load average: 0.04, 0.21, 0.15
[1604222833] SERVICE FLAPPING ALERT: elk;Current Users;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1604222833] SERVICE ALERT: elk;Current Users;OK;SOFT;1;USERS OK - 1 users currently logged in
[1604222868] SERVICE FLAPPING ALERT: elk;PING;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1604222868] SERVICE ALERT: elk;PING;OK;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.13 ms
[1604222981] SERVICE FLAPPING ALERT: elk;Current Users;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1604222981] SERVICE FLAPPING ALERT: elk;PING;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1604223107] SERVICE FLAPPING ALERT: elk;Current Users;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1604223107] SERVICE FLAPPING ALERT: elk;PING;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1604223319] SERVICE FLAPPING ALERT: elk;Current Users;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1604223319] SERVICE FLAPPING ALERT: elk;PING;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1604223446] SERVICE FLAPPING ALERT: elk;Current Users;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1604223446] SERVICE FLAPPING ALERT: elk;PING;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1604223874] SERVICE ALERT: backup;Current Load;OK;HARD;4;OK - load average: 0.01, 0.03, 0.05
[1604223938] SERVICE ALERT: elk;Total Processes;CRITICAL;SOFT;1;CRITICAL - Plugin timed out
[1604223999] SERVICE ALERT: elk;Total Processes;OK;SOFT;2;PROCS OK: 101 processes
[1604224506] SERVICE FLAPPING ALERT: elk;Current Users;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1604224506] SERVICE FLAPPING ALERT: elk;PING;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1604225094] SERVICE ALERT: elk;Swap Usage;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 192.168.1.242: Connection reset by peer
[1604225130] SERVICE ALERT: elk;Zombie Processes;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: Connection refused
[1604225170] SERVICE ALERT: elk;Swap Usage;OK;SOFT;2;SWAP OK - 100% free (2047 MB out of 2047 MB)
[1604225191] SERVICE ALERT: elk;Zombie Processes;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1604227143] SERVICE ALERT: web;Swap Usage;OK;HARD;4;SWAP OK - 99% free (2012 MB out of 2047 MB)
[1604227232] SERVICE ALERT: web;CPU;OK;HARD;4;OK - user: 1.45, nice: 0.50, sys: 4.30, iowait: 0.50, irq: 0.50, softirq: 1.45 idle: 94.78
[1604227237] SERVICE ALERT: web;Zombie Processes;OK;HARD;4;PROCS OK: 0 processes with STATE = Z
[1604227285] SERVICE ALERT: web;Current Users;OK;HARD;4;USERS OK - 0 users currently logged in
[1604227346] SERVICE ALERT: web;Total Processes;OK;HARD;4;PROCS OK: 91 processes
[1604227348] SERVICE ALERT: web;RAM;OK;HARD;4;RAM OK: 59% (140 of 235 MB) used
[1604227380] SERVICE ALERT: web;Root Partition;OK;HARD;4;DISK OK - free space: / 37061 MiB (90.52% inode=100%):
[1604256507] SERVICE ALERT: elk;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604256510] HOST ALERT: elk;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.1.242)
[1604256528] SERVICE ALERT: backup;Current Load;WARNING;SOFT;1;WARNING - load average: 0.05, 0.05, 0.13
[1604256539] SERVICE ALERT: elk;Current Load;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604256543] SERVICE ALERT: web;Current Load;WARNING;SOFT;1;WARNING - load average: 0.02, 0.10, 0.10
[1604256570] SERVICE ALERT: elk;Current Users;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604256570] SERVICE ALERT: elk;CPU;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604256570] HOST ALERT: elk;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.1.242)
[1604256589] SERVICE ALERT: backup;Current Load;WARNING;SOFT;2;WARNING - load average: 0.02, 0.04, 0.13
[1604256602] SERVICE ALERT: elk;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.1.242)
[1604256604] SERVICE ALERT: web;Current Load;WARNING;SOFT;2;WARNING - load average: 0.06, 0.10, 0.10
[1604256633] SERVICE ALERT: elk;RAM;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604256633] HOST ALERT: elk;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.1.242)
[1604256649] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.037 second response time
[1604256649] SERVICE ALERT: backup;Current Load;WARNING;SOFT;3;WARNING - load average: 0.01, 0.04, 0.12
[1604256664] SERVICE ALERT: web;Current Load;WARNING;SOFT;3;WARNING - load average: 0.02, 0.08, 0.09
[1604256665] SERVICE ALERT: elk;Root Partition;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604256696] SERVICE ALERT: elk;SSH;CRITICAL;HARD;1;connect to address 192.168.1.242 and port 22: No route to host
[1604256696] HOST ALERT: elk;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.1.242)
[1604256710] SERVICE ALERT: backup;Current Load;WARNING;HARD;4;WARNING - load average: 0.12, 0.06, 0.13
[1604256725] SERVICE ALERT: web;Current Load;WARNING;HARD;4;WARNING - load average: 0.01, 0.06, 0.08
[1604256728] SERVICE ALERT: elk;Swap Usage;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604256759] HOST ALERT: elk;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.1.242)
[1604256759] SERVICE ALERT: elk;Total Processes;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604256768] SERVICE ALERT: elk;Zombie Processes;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604256822] HOST ALERT: elk;DOWN;SOFT;6;CRITICAL - Host Unreachable (192.168.1.242)
[1604256885] HOST ALERT: elk;DOWN;SOFT;7;CRITICAL - Host Unreachable (192.168.1.242)
[1604256948] HOST ALERT: elk;DOWN;SOFT;8;CRITICAL - Host Unreachable (192.168.1.242)
[1604257011] HOST ALERT: elk;DOWN;SOFT;9;CRITICAL - Host Unreachable (192.168.1.242)
[1604257074] HOST ALERT: elk;DOWN;HARD;10;CRITICAL - Host Unreachable (192.168.1.242)
[1604257325] SERVICE ALERT: web;Current Load;OK;HARD;4;OK - load average: 0.00, 0.02, 0.05
[1604257910] SERVICE ALERT: backup;Current Load;OK;HARD;4;OK - load average: 0.00, 0.01, 0.05
[1604260678] HOST ALERT: elk;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.72 ms
[1604260803] SERVICE ALERT: elk;PING;OK;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.25 ms
[1604260894] SERVICE ALERT: elk;SSH;OK;SOFT;1;SSH OK - OpenSSH_7.4 (protocol 2.0)
[1604261225] SERVICE ALERT: web;Current Load;WARNING;SOFT;1;WARNING - load average: 0.17, 0.06, 0.06
[1604261285] SERVICE ALERT: web;Current Load;OK;SOFT;2;OK - load average: 0.06, 0.05, 0.05
[1604261558] SERVICE ALERT: elk;Total Processes;OK;SOFT;1;PROCS OK: 93 processes
[1604261566] SERVICE ALERT: elk;Zombie Processes;OK;SOFT;1;PROCS OK: 0 processes with STATE = Z
[1604261668] SERVICE ALERT: elk;Current Users;OK;SOFT;1;USERS OK - 1 users currently logged in
[1604261731] SERVICE ALERT: elk;RAM;OK;SOFT;1;RAM OK: 53% (125 of 235 MB) used
[1604261762] SERVICE ALERT: elk;Root Partition;OK;SOFT;1;DISK OK - free space: / 35361 MiB (86.37% inode=100%):
[1604261826] SERVICE ALERT: elk;Swap Usage;OK;SOFT;1;SWAP OK - 95% free (1932 MB out of 2047 MB)
[1604263446] SERVICE ALERT: web;Current Load;WARNING;SOFT;1;WARNING - load average: 0.15, 0.06, 0.06
[1604263506] SERVICE ALERT: web;Current Load;OK;SOFT;2;OK - load average: 0.05, 0.05, 0.05
[1604264768] SERVICE ALERT: web;Current Load;WARNING;SOFT;1;WARNING - load average: 0.16, 0.06, 0.06
[1604264828] SERVICE ALERT: web;Current Load;OK;SOFT;2;OK - load average: 0.06, 0.05, 0.05
[1604265322] SERVICE ALERT: backup;RAM;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.241 port 5666: Network is unreachable
[1604265322] HOST ALERT: backup;DOWN;SOFT;1;CRITICAL - Network Unreachable (192.168.1.241)
[1604265358] HOST ALERT: backup;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.01 ms
[1604265382] SERVICE ALERT: backup;RAM;OK;SOFT;2;RAM OK: 59% (139 of 235 MB) used
[1604265425] SERVICE ALERT: elk;Swap Usage;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: Connection refused
[1604265468] SERVICE ALERT: elk;Zombie Processes;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604265471] HOST ALERT: elk;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.1.242)
[1604265474] SERVICE ALERT: elk;Total Processes;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604265488] SERVICE ALERT: elk;Swap Usage;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604265531] SERVICE ALERT: elk;Zombie Processes;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604265531] HOST ALERT: elk;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.1.242)
[1604265570] SERVICE ALERT: elk;Current Users;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604265594] HOST ALERT: elk;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.1.242)
[1604265602] SERVICE ALERT: elk;PING;CRITICAL;HARD;1;CRITICAL - Host Unreachable (192.168.1.242)
[1604265633] SERVICE ALERT: elk;RAM;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604265657] HOST ALERT: elk;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.1.242)
[1604265665] SERVICE ALERT: elk;Root Partition;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.1.242 port 5666: No route to host
[1604265696] SERVICE ALERT: elk;SSH;CRITICAL;HARD;1;connect to address 192.168.1.242 and port 22: No route to host
[1604265720] HOST ALERT: elk;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.1.242)
[1604265783] HOST ALERT: elk;DOWN;SOFT;6;CRITICAL - Host Unreachable (192.168.1.242)
[1604265846] HOST ALERT: elk;DOWN;SOFT;7;CRITICAL - Host Unreachable (192.168.1.242)
[1604265909] HOST ALERT: elk;DOWN;SOFT;8;CRITICAL - Host Unreachable (192.168.1.242)
[1604265972] HOST ALERT: elk;DOWN;SOFT;9;CRITICAL - Host Unreachable (192.168.1.242)
[1604266035] HOST ALERT: elk;DOWN;HARD;10;CRITICAL - Host Unreachable (192.168.1.242)
[1604268489] SERVICE ALERT: web;Current Load;WARNING;SOFT;1;WARNING - load average: 0.05, 0.10, 0.08
[1604268549] SERVICE ALERT: web;Current Load;WARNING;SOFT;2;WARNING - load average: 0.02, 0.08, 0.08
[1604268610] SERVICE ALERT: web;Current Load;WARNING;SOFT;3;WARNING - load average: 0.01, 0.07, 0.07
[1604268670] SERVICE ALERT: web;Current Load;WARNING;HARD;4;WARNING - load average: 0.00, 0.06, 0.06
[1604268971] SERVICE ALERT: web;Current Load;OK;HARD;4;OK - load average: 0.02, 0.04, 0.05
[1604271530] SERVICE ALERT: elk;Zombie Processes;OK;HARD;4;PROCS OK: 0 processes with STATE = Z
[1604271535] HOST ALERT: elk;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 3.51 ms
[1604271588] SERVICE ALERT: elk;Current Users;OK;HARD;4;USERS OK - 0 users currently logged in
[1604271603] SERVICE ALERT: elk;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 1.40 ms
[1604271665] SERVICE ALERT: elk;Root Partition;OK;HARD;4;DISK OK - free space: / 34356 MiB (83.92% inode=99%):
[1604271681] SERVICE ALERT: elk;RAM;OK;HARD;4;RAM OK: 57% (135 of 235 MB) used
[1604271695] SERVICE ALERT: elk;SSH;OK;HARD;4;SSH OK - OpenSSH_7.4 (protocol 2.0)
[1604271760] SERVICE ALERT: elk;Total Processes;OK;HARD;4;PROCS OK: 101 processes
[1604271787] SERVICE ALERT: elk;Swap Usage;OK;HARD;4;SWAP OK - 93% free (1895 MB out of 2047 MB)